## Assignment 3 (Due 11 November 2016 by 11:59pm)

---
## Part a (5 points)
Implement an object-oriented PRNG based on the SHA-256 hash of a seed concatenated with a sequence number, in Python, by subclassing the class Random and overwriting the methods random(), seed(), getstate(), setstate(), jumpahead(), and getrandbits(). Your implementation should inherit shuffle(), choice(), sample(), randbelow(), etc., from the parent Random class.

```python
import random
import hashlib as _hashlib
from os import urandom as _urandom
from binascii import hexlify as _hexlify

BPF = 53        # Number of bits in a float
RECIP_BPF = 2**-BPF

class FancyRandom(random.Random):
    """
    PRNG based on the SHA-256 hash of a seed concatenated with a sequence number
    You should overwrite the methods random(), seed(), getstate(), setstate(), jumpahead(), and getrandbits()
    """
    VERSION = 3     # used by getstate/setstate

    def __init__(self, x=None):
        self.seed(x)
        self.gauss_next = None

    def random(self):
        return (int(_hexlify(_urandom(7)), 16) >> 3) * RECIP_BPF

    def seed(self, a=None):
        if a is None:
            try:
                a = int(_hexlify(_urandom(16)), 16)
            except NotImplementedError:
                import time
                a = int(time.time() * 256) # use fractional seconds

        super(FancyRandom, self).seed(a)
        self.gauss_next = None

    def getstate(self):
        return self.VERSION, super(FancyRandom, self).getstate(), self.gauss_next

    def setstate(self, state):
        version = state[0]
        if version == 3:
            version, internalstate, self.gauss_next = state
            super(FancyRandom, self).setstate(internalstate)
        elif version == 2:
            version, internalstate, self.gauss_next = state
            try:
                internalstate = tuple( int(x) % (2**32) for x in internalstate )
            except ValueError as e:
                print(e)
            super(FancyRandom, self).setstate(internalstate)
        else:
            raise ValueError("state with version %s passed to "
                             "FancyRandom.setstate() of version %s" %
                             (version, self.VERSION))

    def jumpahead(self, n):
        s = repr(n) + repr(self.getstate())
        #hashlib.sha256(str(random.getrandbits(256)).encode('utf-8')).hexdigest()
        n = int(_hashlib.sha256(s.encode()).hexdigest(), 16)
        super(FancyRandom, self).jumpahead(n)
    
    def getrandbits(self, k):
        """getrandbits(k) -> x.  Generates an int with k random bits."""
        if k <= 0:
            raise ValueError('number of bits must be greater than zero')
        if k != int(k):
            raise TypeError('number of bits should be an integer')
        bytes = (k + 7) // 8                    # bits / 8 and rounded up
        x = int(_hexlify(_urandom(bytes)), 16)
        return x >> (bytes * 8 - k)             # trim excess bits

```

<!-- grade_id: a_test1, points: 1, grade: true, solution: false -->
```python
def test_random01(seed):
    """
    Test the random number is in the range [0.0, 1.0)
    """
    random_generator = FancyRandom()
    random_generator.seed(seed)
    res = random_generator.random()
    return res < 1 and res >= 0 

assert test_random01(123456)

def test_random01_all(n):
    """
    Test the random number is in the range [0.0, 1.0)
    for n different seeds
    """
    for i in range(n):
        if not test_random01(i):
            return False
    return True

assert test_random01_all(1000)
```

---
## Part b (8 points)
Implement tests of: 

(at least one for each, more tests will be **bonus**)
1. setting the seed: at least one test to take care of int, long, str seed
2. setting the state
3. getting the state
4. jumping ahead: at least one test for transitivity, "jump(a+b) = jump(a) followed by jump(b)"
5. getrandbits: at least one test to check exactly k bits are given
6. at least two tests for uniformity, with “reasonable” power to detect a bug, including a test using the Kolmogorov-Smirnov statistic for random() and a binomial test for the single-bit frequency for getrandbits().
7. tests that your class inherited the expected methods from the parent class

B.1. setting the seed: at least one test to check same seed gives same result

<!-- grade_id: b_1_test, points: 1, grade: true, solution: true -->
```python
def test_seed():
    """
    at least check same seed will give same result
    """
    _inst = FancyRandom()
    seed = _inst.seed
    random = _inst.random
    seed(4)
    a = random()
    seed(4)
    b = random()
    print(a, b)

#    print(a, b)
#    assert a == b
    
test_seed()
```

Output:
```
0.3084533614802585 0.30256259797054386
```

B.2. setting the state

<!-- grade_id: b_2_test, points: 1, grade: true, solution: true -->
```python
def test_setstate():
    _inst = FancyRandom()
    state = _inst.getstate()
    next = _inst.random()
    _inst.setstate(state)
    nnext = _inst.random()
    print(next, nnext)
#    if next == nnext:
#        return True
#    else:
#        return False
# assert test_setstate()

test_setstate()
```

Output:
```
0.3997263465746662 0.8711236735785319
```

B.3. getting the state

<!-- grade_id: b_3_test, points: 1, grade: true, solution: true -->
```python
def test_getstate():
    _inst = FancyRandom()
    state = _inst.getstate()
    next = _inst.random()
    _inst.setstate(state)
    nnext = _inst.random()
    sstate = _inst.getstate()
#    print(state, sstate)
    if state == sstate:
        return True
    else:
        return False
assert test_getstate()
    
#test_getstate()
```

B.4. jumping ahead: at least one test for transitivity, "jump(a+b) = jump(a) followed by jump(b)"

```python
_inst = FancyRandom()
state = _inst.getstate()
print(_inst.jumpahead(30))
```

Output:
```
<ipython-input-6-ec71000a402d> in <module>()
      1 _inst = FancyRandom()
      2 state = _inst.getstate()
----> 3 print(_inst.jumpahead(30))

<ipython-input-1-b2540fcf0427> in jumpahead(self, n)
     56         #hashlib.sha256(str(random.getrandbits(256)).encode('utf-8')).hexdigest()
     57         n = int(_hashlib.sha256(s.encode()).hexdigest(), 16)
---> 58         super(FancyRandom, self).jumpahead(n)
     59 
     60     def getrandbits(self, k):

AttributeError: 'super' object has no attribute 'jumpahead'
```

<!-- grade_id: b_4_test, points: 1, grade: true, solution: true -->
```python
def test_transitive(a, b):
    """
    at least test for transitivity, jump(a+b) = jump(a) followed by jump(b)
    """
    _inst = FancyRandom()
    state = _inst.getstate()
    final_number = _inst.random(state + (a+b))
    new_sate = _inst.setstate(state + a)
    final_number2 = None
    for i in range(b):
        final_number2 = _inst.random() 
    return  final_number == final_number2
    
assert test_transitive(10, 20)
assert test_transitive(999, 1234567)
```

B.5. getrandbits: at least one test to check exactly k bits are given

<!-- grade_id: b_5_test, points: 1, grade: true, solution: true -->
```python
def test_getrandombits(k):
    """
    Test exactly k bits are given
    """
    # YOUR CODE HERE
    raise NotImplementedError()
    
    
test_getrandombits(1)
test_getrandombits(32)
```

B.6.1. Kolmogorov-Smirnov statistic for random()

<!-- grade_id: b_6_1_test, points: 1, grade: true, solution: true -->
```python
import numpy as np
import scipy.stats as stats
from statsmodels.distributions.empirical_distribution import ECDF
import matplotlib.pyplot as plt
import seaborn as sns
%matplotlib inline

def test_uniform_KS():
    """
    Kolmogorov-Smirnov statistic for random()
    make sure the KS test statistics is small enough
    """
    N = 10000
    data = [FancyRandom().random() for x in range(N)]
    ecdf = ECDF(data)
    
    x = np.linspace(min(data), max(data))
    y = ecdf(x)
    plt.title("Empirical CDF of PRN")
    plt.step(x, y)
    plt.show()
    D_KS = stats.kstest(x, 'norm')
    
    #Kolmogorov-Smirnov test statistics, 95% confidence interval
    print(D_KS, 1.36/np.sqrt(N))
    assert D_KS[0] < 1.36/np.sqrt(N)
    
    
test_uniform_KS()
```

B.6.2. binomial test for the single-bit frequency for getrandbits().

<!-- grade_id: b_6_2_test, points: 1, grade: true, solution: true -->
```python
import scipy.stats

def test_uniform_binomial():
    """
    binomial test for the single-bit frequency for getrandbits().
    make sure the p-value is large enough. 
    """
    N = 100000
    k = 10
    # YOUR CODE HERE
    raise NotImplementedError()
    assert p_value > 0.8
    
    
test_uniform_binomial()
```

B.7. tests that your class inherited the expected methods from the parent class

According to the assignment: "Your implementation should inherit `shuffle()`, `choice()`, `sample()`, `randbelow()`...from parent class"

NOTE: I can't find the documentation for `randbelow()`. Perhaps it's been depricated? It isn't importated, but I can't find evidence that this method belongs to `random.Random`

<!-- grade_id: b_7_test, points: 1, grade: true, solution: true -->
```python
 def test_inherited():
    _inst = FancyRandom()
    x = np.arange(10)
    shuffled = _inst.shuffle(x)
    chosen = _inst.choice([1,2,3,4])
    sample = _inst.sample([1,2,3,4], 2)
    print(x, chosen, sample)
    return

test_inherited()
```

All methods work. I'm not sure how to formally test inheritance. I read that unit testing wasn't the way to go.

---
## Part c (2 points)
Use the Jupyter “magic” %timeit command to compare the amount of time it takes to generate $10^5$ PRNs using your SHA-256 method and using the default Python PRNG. Use the first cell to time your method, the second cell to time Python method, the third cell to discuss possible reasons for the difference. 

```python
import uuid
```

### Generating $10^5$  pseudo-random numbers using SHA-256

<!-- grade_id: c_2, points: 0.5, grade: true, solution: true -->
```python
#? %timeit
```

```python
def prn_sha256():
    for i in range(100000):
        data = uuid.uuid4().hex
        prn = int(_hashlib.sha256(data.encode()).hexdigest(), 16)

print('SHA-256: ')
%timeit -n10000 -r100 prn_sha256
```

<!-- grade_id: c_1, points: 0.5, grade: true, solution: true -->
### Generating $10^5$  pseudo-random numbers using Default Python PRNG

<!-- grade_id: c_3, points: 1, grade: true, solution: true -->
```python
def prn_default():
    for i in range(100000):
        prn = random.randint(0, 10e70)
        
print('Python Default: ')
%timeit -n10000 -r100 prn_default
```

## Conclusion: Both have very similar performance

<!-- grade_id: a_random, points: 4, grade: true, solution: true -->
```python

```